<a href="https://colab.research.google.com/github/TheChroniclerr/CP322/blob/main/Assignment_3.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# <center><u> CP322: Assignment 3 – Winter 2023 </u> </center>
# <center><u>  Due on  March 01, 2026 (Before 11:59 PM) </u> </center>

 This is a **group (of 2)** assignment, and we will practice the advancing concepts to build classification and regression models etc. There are two parts to this assignment.

 For this assignment, you must use Python language. You will use this jupyter notebook to write your code without errors. If your code will not run, then you will score zero. Therefore, ensure you have removed all syntax errors from your code and not changing the variable names. Gradescope platform would be used to upload the assignments for grading and autograded question's feedback will be visible to you. The link to the Gradescope assignment is available on Myls course page.

 For submission, drag and drop your code file(s) into Gradescope and add your partner's name. Make sure that your file name should be as suggested in the assignment, using a different name may score zero.

 - Please note that the submitted code will be checked for plagiarism. By submitting this zip file, you would confirm that you have not received unauthorized assistance in preparing the assignment. You also confirm that you are aware of course policies for submitted work.

 - Marks will be deducted from any questions where these requirements are not met.

 - Multiple attempts will be allowed, and only your last submission before the deadline will be graded. Instructor reserves the right to take off points for not following the directions.

<b>Warning:</b> Follow the assignment instructions to the letter in terms of the variable names and function names, as this assignment will be auto-graded except few written type questions.

**Important Points**:

- For each question in the assignment, please write down your answer in the answer cell(s) right below the question.
- It is helpful to have extra cells breaking down the process towards reaching your final answer. If you happen to create new cells below your answer to run codes, **NEVER** add cells between a question cell and the answer cell below it. It will cause errors when autograder run.
- Finally, unless it is stated otherwise, try to avoid using python for loops or list comprehensions.  The majority of this part of assignment can be done using builtin commands in Pandas, numpy, seaborn, matplotlib, and sklearn functions.  
- **Partners.** You may work with one other partner; **Only one person is required to submit to the assignment. On Gradescope, the person who submits *must* also designate their partner so that both of you receive credit.** Once you submit, click into your submission, and there will be an option to Add Group Member in the top right corner. Make sure to **re-add your partner every time you make a new submission**, as Gradescope does not save partner information.

- **Rules.** Don't share your code with anybody but your partner. You are welcome to discuss questions with other students, but don't share the answers. The experience of solving the problems in this project will prepare you for exams (and life). If someone asks you for the answer, resist! Instead, you can demonstrate how you would solve a similar problem.

- **Tests.** The tests that are given are **not comprehensive** and passing the tests for a question **does not** mean that you answered the question correctly. However, more tests will be applied to verify the correctness of your submission in order to assign your final score, so be careful and check your work! You might want to create your own checks along the way to see if your answers make sense. Additionally, before you submit, make sure that none of your cells take a very long time to run (several minutes). Always it is good to re-run your assignment cells before submission.

- **Free Response/written Questions.** Make sure that you put the answers to the written questions in the indicated cell we provide. **Every free response question should include an explanation** that adequately answers the question. The written questions will be manually graded on Gradescope.

**Advice.** Develop your answers incrementally. To perform a complicated task, break it up into steps, perform each step on a different line, give a new name to each result, and check that each intermediate result is what you expect. You can add any additional names or functions you want to the provided cells. Make sure that you are using distinct and meaningful variable names throughout the notebook. Along that line, **DO NOT** reuse the variable names that we use when we grade your answers.

To get started, load the required packages. Make sure to also run the first cell of this notebook to load `otter`.

# Part 1: Handmade Classifier

 In this part, you will fit a classification model that guesses whether a movie is a comedy or a thriller by using only the number of times chosen words appear in the movie's screenplay. you will fit a class of classification algorithm to make predictions.


```python
# Run this cell to suppress all FutureWarnings
import warnings
warnings.filterwarnings("ignore", category=FutureWarning)

# for data manipulation
import pandas as pd, numpy as np, random

#for graphs
import matplotlib.pyplot as plt
import seaborn as sns
%matplotlib inline

#models to run
from sklearn.linear_model import LinearRegression
from sklearn.tree import DecisionTreeClassifier
from sklearn import preprocessing
from sklearn import neighbors
from sklearn import tree

#train_test_split
from sklearn.model_selection import train_test_split,StratifiedShuffleSplit, GridSearchCV

#for cycle
from itertools import cycle

#metrics
from sklearn import metrics
from sklearn.metrics import confusion_matrix, classification_report, precision_score, recall_score, f1_score, r2_score, accuracy_score, roc_auc_score
from sklearn.model_selection import cross_val_score, KFold
from sklearn.ensemble import RandomForestClassifier
from sklearn.neighbors import KNeighborsClassifier
from sklearn.model_selection import train_test_split, GridSearchCV

RANDOM_STATE = 42
random.seed(RANDOM_STATE)

import otter
grader = otter.Notebook()
```

### Dataset Loading

Here, we are exploring movie screenplays. We'll be trying to predict each movie's genre from the text of its screenplay. In particular, we have compiled a list of 5,000 words that occur in conversations between movie characters. For each movie, our dataset tells us the frequency with which each of these words occurs in certain conversations in its screenplay. All words have been converted to lowercase.

Run the cell below to read the `movies` table. **It may take up to a minute to load.**

```python
movies = pd.read_csv('https://raw.githubusercontent.com/sukhjitsehra/datasets/master/CP322/movies.csv')
movies
```

Here is one row of the table and some of the frequencies of words that were said in the movie.

```python
movies.iloc[movies.index[movies["Title"]=="runaway bride"],[0, 1, 2, 3, 4, 14, 49, 1042, 4004]]
```

The above cell prints a few columns of the row for the comedy movie *Runaway Bride*.  The movie contains 4895 words. The word "it" appears 115 times, as it makes up  $\frac{115}{4895} \approx 0.0234092$ of the words in the movie. The word "england" doesn't appear at all.

Additional context: This numerical representation of a body of text, one that describes only the frequencies of individual words, is called a bag-of-words representation. This is a model that is often used in [NLP](https://en.wikipedia.org/wiki/Natural_language_processing). A lot of information is discarded in this representation: the order of the words, the context of each word, who said what, the cast of characters and actors, etc. However, a bag-of-words representation is often used for machine learning applications as a reasonable starting point, because a great deal of information is also retained and expressed in a convenient and compact format.

Now, we will investigate whether this representation is sufficient to build an accurate genre classifier.

All movie titles are unique. The `row_for_title` function provides fast access to the one row for each title.

*Note: All movies in our dataset have their titles lower-cased.*

```python
def row_for_title(title):
    # Find the row where the column "Title" matches the input title
    row = movies.loc[movies["Title"] == title]

    # Return the index of the row
    return movies.iloc[row.index[0],:]
row_for_title('toy story')
```

For example, the fastest way to find the frequency of "fun" in the movie *Toy Story* is to access the `'fun'` item from its row. Check the original table to see if this worked for you!

```python
row_for_title('toy story')['fun']
```

**Question 1.0**

Set `expected_row_sum` to the number that you __expect__ will result from summing all proportions in each row, excluding the first five columns. Think about what any one row adds up to.


```python
# Set row_sum to a number that's the (approximate) sum of each row of word proportions.
expected_row_sum = ...
```

This dataset was extracted from [a dataset from Cornell University](http://www.cs.cornell.edu/~cristian/Cornell_Movie-Dialogs_Corpus.html). After transforming the dataset (e.g., converting the words to lowercase, removing the naughty words, and converting the counts to frequencies), we created this new dataset containing the frequency of 5000 common words in each movie.

```python
print('Words with frequencies:', len(movies.axes[1]) - 5) # except first 5
print('Movies with genres:', len(movies.axes[0]))
```

#### 1.1. Word Stemming
The columns other than "Title", "Year", "Rating", "Genre", and "# Words" in the `movies` table are all words that appear in some of the movies in our dataset.  These words have been *stemmed*, or abbreviated heuristically, in an attempt to make different [inflected](https://en.wikipedia.org/wiki/Inflection) forms of the same base word into the same string.  For example, the column "manag" is the sum of proportions of the words "manage", "manager", "managed", and "managerial" (and perhaps others) in each movie. This is a common technique used in machine learning and natural language processing.

Stemming makes it a little tricky to search for the words you want to use, so we have provided another table called `vocab_table` that will let you see examples of unstemmed versions of each stemmed word. Run the code below to load it.

**Note:** You should use `vocab_table` for the rest of Section, not `vocab_mapping`.

```python
# Just run this cell.
vocab_mapping = pd.read_csv('https://raw.githubusercontent.com/sukhjitsehra/datasets/master/CP322/stem.csv')
stemmed = np.take(movies.columns, np.arange(3, len(movies.columns)))
vocab_table = pd.merge(pd.DataFrame({'Stem': stemmed}), vocab_mapping, on='Stem')
vocab_table.iloc[1100:1110]
```

**Question 1.1.1**

Using `vocab_table`, find the stemmed version of the word "elements" and assign the value to `stemmed_message`.

```python
stemmed_message = ...
stemmed_message
```

**Question 1.1.2**

What stem in the `vocab_table` dataset has the most words that are shortened to it? Assign `most_stem` to that stem.

```python
most_stem = ...
most_stem
```

**Question 1.1.3**

What is the longest word in the dataset whose stem wasn't shortened? Assign that to `longest_uncut`. Break ties alphabetically from Z to A (so if your options are "cat" or "bat", you should pick "cat"). Note that when sorting letters, the letter `a` is smaller than the letter `z`.

*Hint:* `vocab_table` has 2 columns: one for stems and one for the unstemmed (normal) word. Find the longest word that wasn't cut at all (same length as stem).

```python
# Create a new DataFrame with two new columns: "Word len" and "Stem len"
vocab_table['Word len'] = vocab_table['Word'].apply(len)
vocab_table['Stem len'] = vocab_table['Stem'].apply(len)

# Create a new column "len dif" that shows the difference between "Word len" and "Stem len"
vocab_table['len dif'] = ...

# Filter the DataFrame to only include rows where "len dif" is 0
tbl_with_diff = ...

# Sort the filtered DataFrame by "Word" in descending order and "Word len" in descending order
tbl_with_diff = ...

# Extract the first value in the "Word" column
longest_uncut = ...
longest_uncut
```

#### 1.2. Exploratory Data Analysis: Linear Regression

Let's explore our dataset before trying to build a classifier. To start, we’ll use the associated proportions to investigate the relationship between different words.

The first association we'll investigate is the association between the proportion of words that are "outer" and the proportion of words that are "space".

As usual, we'll investigate our data visually before performing any numerical analysis.

Run the cell below to plot a scatter diagram of "space" proportions vs "outer" proportions and to create the `outer_space` dataframe. Each point on the scatter plot represents one movie.

```python
# Select the "outer" and "space" columns from the DataFrame
outer_space = movies[['outer', 'space']]

# Create a scatter plot of "outer" proportions vs "space" proportions
outer_space.plot.scatter(x='outer', y='space', figsize=(10, 8))
plt.xlim(-0.0005, 0.001)
plt.ylim(-0.0005, 0.003)
# Rotate the x-axis labels by 45 degrees
plt.xticks(rotation=45)
plt.draw()
```

**Question 1.2.1**

Looking at that chart it is difficult to see if there is an association. Calculate the correlation coefficient for the potential linear association between proportion of words that are "outer" and the proportion of words that are "space" for every movie in the dataset, and assign it to `outer_space_cor`.

```python
# These two arrays should make your code cleaner!
outer = movies["outer"]
space = movies["space"]

# code to find (X_i - mean(X))/std(X)
outer_su = ...
space_su = ...

# find correlation
outer_space_cor = ...
outer_space_cor
```

**Question 1.2.2 (Manually Graded)**

Choose two *different* words in the movies dataset with a magnitude (absolute value) of correlation higher than 0.2 and plot a scatter plot with a line of best fit for them. Please do not pick "outer" and "space" or "san" and "francisco". The code to plot the scatter plot and line of best fit is given for you, you just need to calculate the correct values to `cor`, `slope` and `intercept`.

*Hint:* It's easier to think of words with a positive correlation, i.e. words that are often mentioned together*. Try to think of common phrases or idioms.

<!-- BEGIN QUESTION -->



```python
word_x = "cold"
word_y = "warm"

# These arrays should make your code cleaner!
arr_x = movies["cold"]
arr_y = movies["warm"]

# code to find (X_i - mean(X))/std(X)
x_su = ...
y_su = ...

cor = ...

slope = ...
intercept = ...

# DON'T CHANGE THESE LINES OF CODE
# Create a scatter plot of the data
ax = movies.plot.scatter(x=word_x, y=word_y, figsize=(10, 8), color='blue')

# Add the regression line to the plot
max_x = arr_x.max()
ax.plot([0, max_x * 1.3], [intercept, intercept + slope * (max_x*1.3)], color='gold', label='Regression Line')

# Add the title to the plot
plt.title(f"Correlation: {cor:.3f}, magnitude greater than .2: {abs(cor) >= 0.2}")

# Show the plot
plt.draw()
```

<!-- END QUESTION -->

**Question 1.2.3**

Imagine that you picked the words "san" and "francisco" as the two words that you would expect to be correlated because they compose the city name San Francisco. Assign `san_francisco` to either the number '1' or  '2' according to which statement is true regarding the correlation between "san" and "francisco."

1. "san" can also preceed other city names like San Diego and San Jose. This might lead to "san" appearing in movies without "francisco," and would reduce the correlation between "san" and "francisco."
2. "san" can also preceed other city names like San Diego and San Jose. The fact that "san" could appear more often in front of different cities and without "francisco" would increase the correlation between "san" and "francisco."

```python
san_francisco = ...
```

#### 1.3. Splitting the dataset

Now, we're going to use our `movies` dataset for two purposes.

1. First, we want to *train* movie genre classifiers.
2. Second, we want to *test* the performance of our classifiers.

Hence, we need two different datasets: *training* and *test*.

The purpose of a classifier is to classify unseen data that is similar to the training data. The test dataset will help us determine the accuracy of our predictions by comparing the actual genres of the movies with the genres that our classifier predicts. Therefore, we must ensure that there are no movies that appear in both sets. We do so by splitting the dataset randomly. The dataset has already been permuted randomly, so it's easy to split.  We just take the first 85% of the dataset for training and the rest for test.

Run the code below (without changing it) to separate the datasets into two tables.

```python
movies.shape[0]
```

```python
training_proportion = 17/20

num_movies = movies.shape[0]
num_train = int(num_movies * training_proportion)
num_test = num_movies - num_train

train_movies = movies.take(np.arange(num_train))
test_movies = movies.take(np.arange(num_train, num_movies))

print("Training: ",   train_movies.shape[0], ";",
      "Test: ",       test_movies.shape[0])
```

**Question 1.3.1 (Manually Graded)**

Draw a horizontal bar chart with two bars that show the proportion of Comedy movies in each dataset (`train_movies` and `test_movies`). The two bars should be labeled "Training" and "Test". Complete the function `comedy_proportion` first; it should help you create the bar chart.

<!-- BEGIN QUESTION -->



```python

def comedy_proportion(table):
    # Return the proportion of movies in a table that have the comedy genre.
    ...
    return ...

# The staff solution took multiple lines.  Start by creating a table.
# If you get stuck, think about what sort of table you need for barh to work
...
```

<!-- END QUESTION -->

### 2 K-Nearest Neighbors

We are going to `simulate` the steps that `K-Nearest Neighbors (k-NN)` algorithm will use for classification, i.e.,  Given some numerical *attributes* (also called *features*) of an unseen example, it decides which category that example belongs to based on its similarity to previously seen examples. Predicting the category of an example is called *labeling*, and the predicted category is also called a *label*.

An attribute (feature) we have about each movie is *the proportion of times a particular word appears in the movie*, and the labels are two movie genres: comedy and thriller.  The algorithm requires many previously seen examples for which both the attributes and labels are known: that's the `train_movies` table.

To build understanding, we're going to visualize the algorithm instead of just describing it.

#### 2.1. Classifying a movie

In k-NN, we classify a movie by finding the `k` movies in the *training set* that are most similar according to the features we choose. We call those movies with similar features the *nearest neighbors*.  The k-NN algorithm assigns the movie to the most common category among its `k` nearest neighbors.

Let's limit ourselves to just 2 features for now, so we can plot each movie.  The features we will use are the proportions of the words "water" and "feel" in the movie.  Taking the movie *Monty Python and the Holy Grail* (in the test set), 0.000804074 of its words are "water" and 0.0010721 are "feel". This movie appears in the test set, so let's imagine that we don't yet know its genre.

First, we need to make our notion of similarity more precise.  We will say that the *distance* between two movies is the straight-line distance between them when we plot their features on a scatter diagram.

**This distance is called the Euclidean ("yoo-KLID-ee-un") distance, whose formula is $\sqrt{(x_1 - x_2)^2 + (y_1 - y_2)^2}$.**

For example, in the movie *Clerks.* (in the training set), 0.00016293 of all the words in the movie are "water" and 0.00154786 are "feel".  Its distance from *Monty Python and the Holy Grail* on this 2-word feature set is $\sqrt{(0.000804074 - 0.000162933)^2 + (0.0010721 - 0.00154786)^2} \approx 0.000798379$.  (If we included more or different features, the distance could be different.)

A third movie, *The Godfather* (in the training set), has 0 "water" and 0.00015122 "feel".

The function below creates a plot to display the "water" and "feel" features of a test movie and some training movies. As you can see in the result, *Monty Python and the Holy Grail* is more similar to *Clerks.* than to the *The Godfather* based on these features, which makes sense as both movies are comedy movies, while *The Godfather* is a thriller.

```python
# Just run this cell.
def plot_with_two_features(test_movie, training_movies, x_feature, y_feature):
    """Plot a test movie and training movies using two features."""
    test_row = row_for_title(test_movie)
    distances = pd.DataFrame({
            x_feature: [test_row[x_feature]],
            y_feature: [test_row[y_feature]],
            'Color': 0,
            'Title': [test_movie]
        })

    color_map = {'unknown': 0, 'comedy': 1, 'drama': 2, 'action': 3, 'thriller': 4}
    for movie in training_movies:
        row = row_for_title(movie)
        distances = pd.concat([distances, pd.DataFrame([{
            x_feature: row[x_feature],
            y_feature: row[y_feature],
            'Color': color_map.get(row['Genre']),
            'Title': movie
        }])], ignore_index=True)
    fig = plt.figure(figsize=(12, 8))
    plt.axis([-0.0008, 0.001, -0.004, 0.007]);
    ax = plt.scatter(distances[x_feature], distances[y_feature], c=distances['Color'], s=50)
    plt.xlabel(x_feature)
    plt.ylabel(y_feature)
    for i, title in enumerate(distances['Title']):
        ax.axes.annotate(title, (distances[x_feature][i], distances[y_feature][i]))
    plt.draw()

training = ["clerks.", "the godfather"]
plot_with_two_features("monty python and the holy grail", training, "water", "feel")
```

**Question 2.1.1**

Compute the Euclidean distance (defined in the section above) between the two movies, *Monty Python and the Holy Grail* and *The Godfather*, using the `water` and `feel` features only.  Assign it the name `one_distance`.

*Hint:* In the formula for Euclidean distance, think carefully about what `x` and `y` represent. Refer to the example in the text above if you are unsure.

```python
python = row_for_title("monty python and the holy grail")
godfather = row_for_title("the godfather")

one_distance = ...
one_distance
```

Below, we've added a third training movie, *The Silence of the Lambs*. Before, the point closest to *Monty Python and the Holy Grail* was *Clerks.*, a comedy movie. However, now the closest point is *The Silence of the Lambs*, a thriller movie.

```python
training = ["clerks.", "the godfather", "the silence of the lambs"]
plot_with_two_features("monty python and the holy grail", training, "water", "feel")
```

**Question 2.1.2**

Complete the function `distance_two_features` that computes the Euclidean distance between any two movies, using two features. The last two lines call your function to show that *Monty Python and the Holy Grail* is closer to *The Silence of the Lambs* than it is to *Clerks*.

```python
def distance_two_features(title0, title1, x_feature, y_feature):
    """Compute the distance between two movies with titles title0 and title1.

    Only the features named x_feature and y_feature are used when computing the distance.
    """
    row0 = ...
    row1 = ...
    ...

for movie in ["clerks.", "the silence of the lambs"]:
    movie_distance = distance_two_features(movie, "monty python and the holy grail", "water", "feel")
    print(movie, 'distance:\t', movie_distance)
```

**Question 2.1.3**

Define the function `distance_from_python` so that it works as described in its documentation.

**Note:** Your solution should not use arithmetic operations directly. Instead, it should make use of existing functionality above!

```python
def distance_from_python(title):
    """The distance between the given movie and "monty python and the holy grail",
    based on the features "water" and "feel".

    This function takes a single argument:
      title: A string, the name of a movie.
    """

    ...

# Calculate the distance between "Clerks." and "Monty Python and the Holy Grail"
distance_from_python('clerks.')
```

**Question 2.1.4**

Using the features `"water"` and `"feel"`, what are the names and genres of the 5 movies in the **training set** closest to *Monty Python and the Holy Grail*?  To answer this question, make a **dataframe** named `close_movies` containing those 5 movies with columns `"Title"`, `"Genre"`, `"water"`, and `"feel"`, as well as a column called `"distance from python"` that contains the distance from *Monty Python and the Holy Grail*.  The table should be **sorted in ascending order by `"distance from python"`**.

*Note:* Why are smaller distances from *Monty Python and the Holy Grail* more helpful in helping us classify the movie?

*Hint:* Your final dataframe should only have 5 rows. How can you get the first five rows of a dataframe?

```python

# The staff solution took multiple lines.
...
close_movies = ...
close_movies
```

**Question 2.1.5**

Next, we'll clasify *Monty Python and the Holy Grail* based on the genres of the closest movies.

To do so, define the function `most_common` so that it works as described in its documentation below.

```python
def most_common(df, column):
    """The most common element in a column of a table.

    This function takes two arguments:
      column: The name of a column, a string.
      df: dataframe

    It returns the most common value in the label column of the table.
    In case of a tie, it returns any one of the most common values.
    """
    ...

# Calling most_common on your table of 5 nearest neighbors classifies
# "monty python and the holy grail" as a thriller movie, 3 votes to 2.
most_common(close_movies, 'Genre')
```

Congratulations are in order -- you've classified your first movie! However, we can see that the classifier doesn't work too well since it categorized *Monty Python and the Holy Grail* as a thriller movie (unless you count the thrilling holy hand grenade scene). Let's see if we can do better!

### 3 Add More Features

Now, we're going to extend our classifier to consider more than two features at a time.

Euclidean distance still makes sense with more than two features. For `n` different features, we compute the difference between corresponding feature values for two movies, square each of the `n`  differences, sum up the resulting numbers, and take the square root of the sum.

**Question 3.0**

Write a function called `distance` to compute the Euclidean distance between two **arrays** of **numerical** features (e.g. arrays of the proportions of times that different words appear). The function should be able to calculate the Euclidean distance between two arrays of arbitrary (but equal) length.

Next, use the function you just defined to compute the distance **between the first and second movie** in the **training set** *using all of the features*.  (Remember that the first five columns of your tables are not features.)

*Hint:* Make sure to skip the first five columns of the dataframe before you compute `distance_first_to_second`, as these columns do not contain any features (the proportions of words).

```python
def distance(features_array1, features_array2):
    """The Euclidean distance between two arrays of feature values."""
    ...

distance_first_to_second = ...
distance_first_to_second
```

#### 3.1. Creating your own feature set

Unfortunately, using all of the features has some downsides.  One clear downside is the lack of *computational efficiency* -- computing Euclidean distances just takes a long time when we have lots of features.  You might have noticed that in the last question!

So in this part, we're going to start with just 10-20.  We'd like to choose features that are very *discriminative*. That is, features which lead us to correctly classify as much of the test set as possible.  This process of choosing features that will make a classifier work well is sometimes called *feature selection*, or, more broadly, *feature engineering*.

In this question, we will help you get started on selecting more effective features for distinguishing comedy from thriller movies. The plot below (generated for you) shows the average number of times each word occurs in a comedy movie on the horizontal axis and the average number of times it occurs in an thriller movie on the vertical axis.


*Note: The line graphed is the line of best fit, NOT the line y=x.*

![word_plot.png](https://github.com/sukhjitsehra/datasets/blob/master/CP322/images/word_plot.png?raw=true)

Questions 3.1.1 through 3.1.4 will ask you to interpret the plot above. For each question, select one of the following choices and assign its number to the provided name. 1. The word is common in both comedy and thriller movies 2. The word is uncommon in comedy movies and common in thriller movies 3. The word is common in comedy movies and uncommon in thriller movies 4. The word is uncommon in both comedy and thriller movies 5. It is not possible to say from the plot

**Question 3.1.1**

What properties do words in the bottom left corner of the plot have? Your answer should be a single integer from 1 to 5, corresponding to the correct statement from the choices above.

```python
bottom_left = ...
```

**Question 3.1.2**

What properties do words in the bottom right corner have? Your answer should be a single integer from 1 to 5, corresponding to the correct statement from the choices above.

```python
bottom_right = ...
```

**Question 3.1.3**

What properties do words in the top right corner have? Your answer should be a single integer from 1 to 5, corresponding to the correct statement from the choices above.

```python
top_right = ...
```

**Question 3.1.4**

What properties do words in the top left corner have? Your answer should be a single integer from 1 to 5, corresponding to the correct statement from the choices above.

```python
top_left = ...
```

**Question 3.1.5**

If we see a movie with a lot of words that are common for comedy movies but uncommon for thriller movies, what would be a reasonable guess about the genre of the movie? Assign `movie_genre` to the integer corresponding to your answer:
1. if it is a thriller movie.
2. if it is a comedy movie.

```python
movie_genre_guess = ...
```

**Question 3.1.6**

Using the plot above, make an list of at least 10 common words that you think might let you **distinguish** between comedy and thriller movies. Make sure to choose words that are **frequent enough** that every movie contains at least one of them. Don't just choose the most frequent words though--you can do much better.

```python
# Set my_features to an list (strings that are column labels) e.g. 'i', 'the', 'to', 'a', 'it', 'and', 'that', 'of', 'your', 'what', 'in', 'me', 'is', 'do', 'thi', 'dont', 'he', 'for', 'know'

my_features = ...

# Select the above features from both the train and test sets
train_my_features = ...
test_my_features = ...

```

This test makes sure that you have chosen words such that at least one appears in each movie. If you can't find words that satisfy this test just through intuition, try writing code to print out the titles of movies that do not contain any words from your list, then look at the words they do contain.

Next, let's classify the first movie from our test set using these features.  You can examine the movie by running the cells below. Do you think it will be classified correctly?

```python
print("Movie:")
test_movies[['Title', 'Genre']].iloc[0]
```

```python
print("Features:")
test_my_features.iloc[0]
```

As before, we want to look for the movies in the training set that are most like our test movie.  We will calculate the Euclidean distances from the test movie (using `my_features`) to all movies in the training set.  You could do this with a `for` loop, but to make it computationally faster, we have provided a function, `fast_distances`, to do this for you.  Read its documentation to make sure you understand what it does.  (You don't need to understand the code in its body unless you want to.)

```python
def fast_distances(test_row, train_df):
    assert train_df.shape[1] < 50, "Make sure you're not using all the features of the movies table."
    assert type(test_row) != pd.DataFrame, "Make sure you are passing in a row object to fast_distances."
    assert len(test_row) == train_df.shape[1], "Make sure the length of test row is the same as the number of columns in train_table."
    counts_matrix = train_df.values
    diff = np.tile(np.array(test_row), [counts_matrix.shape[0], 1]) - counts_matrix
    np.random.seed(0) # For tie breaking purposes
    distances = np.squeeze(np.asarray(np.sqrt(np.square(diff).sum(1))))
    eps = np.random.uniform(size=distances.shape)*1e-10 #Noise for tie break
    distances = distances + eps
    return distances
```

**Question 3.1.7**

Use the `fast_distances` function provided above to compute the distance from the first movie in your test set to all the movies in your training set, **using your set of features**. Make a new dataframe called `genre_and_distances` with one row for each movie in the training set and two columns:
* The `"Genre"` of the training movie
* The `"Distance"` from the first movie in the test set

Ensure that `genre_and_distances` is **sorted in ascending order by distance to the first test movie**.

*Hint:* Think about how you can use the variables you defined in 3.1.6.

```python
# The staff solution took multiple lines of code.
genre_and_distances = ...
print(genre_and_distances.iloc[0:4])
```

**Question 3.1.8**

Now compute the 7-nearest neighbors classification of the first movie in the test set.  That is, decide on its genre by finding the most common genre among its 7 nearest neighbors in the training set, according to the distances you've calculated.  Then check whether your classifier chose the right genre.  (Depending on the features you chose, your classifier might not get this movie right, and that's okay.)

*Hint:* You should use the `most_common` function that was defined earlier and pass first 7 rows of data frame `genre_and_distances`.

```python
# Set my_assigned_genre to the most common genre among these.
my_assigned_genre = ...

# Set my_assigned_genre_was_correct to True if my_assigned_genre
# matches the actual genre of the first movie in the test set, False otherwise.
my_assigned_genre_was_correct = ...

print("The assigned genre, {}, was{}correct.".format(my_assigned_genre, " " if my_assigned_genre_was_correct else " not "))
```

#### 3.2. A classifier function

Now we can write a single function that encapsulates the whole process of classification.

**Question 3.2.1**

Write a function called `classify`.  It should take the following four arguments:
* A row of features for a movie to classify (e.g., `test_my_features.iloc[0]`).
* A dataframe with a column for each feature (e.g., `train_my_features`).
* An array of classes (e.g. the labels "comedy" or "thriller") that has as many items as the previous table has rows, and in the same order. *Hint:* What are the labels of each row in the training set?
* `k`, the number of neighbors to use in classification.

It should return the class label that a `k`-nearest neighbor classifier picks for the given row of features (the string `'comedy'` or the string `'thriller'`).

```python
def classify(test_row, train_features, train_labels, k):
    """Return the most common class among k nearest neigbors to test_row."""
    distances = fast_distances(test_row, train_features)
    genre_and_distances = ...
    ...
```

**Question 3.2.2**

Assign `godzilla_genre` to the genre predicted by your classifier for the movie "godzilla" in the test set, using **15 neighbors** and using your 10 features.

*Hint:* The `row_for_title` function will not work here.

```python
# The staff solution first defined a row called godzilla_features.
godzilla_features = ...
godzilla_genre = ...
godzilla_genre
```

Finally, when we evaluate our classifier, it will be useful to have a classification function that is specialized to use a fixed training set and a fixed value of `k`.

**Question 3.2.3**

Create a classification function that takes as its argument a row containing your 10 features and classifies that row using the 10-nearest neighbors algorithm with `train_my_features` as its training set.

```python
def classify_feature_row(row):
    ...

# When you're done, this should produce 'thriller' or 'comedy'.
classify_feature_row(test_my_features.iloc[0])
```

#### 3.3. Evaluating your classifier

**Question 3.3.1**

Use `classify_feature_row` and `apply` to classify every movie in the test set.  Assign these guesses as an array to `test_guesses`.  Then, compute the proportion of correct classifications.

*Hint*: Which dataset do you want to apply this function to?

```python
test_guesses = ...
proportion_correct = ...
proportion_correct
```

**Question 3.3.2**

An important part of evaluating your classifiers is figuring out where they make mistakes. Assign the name `test_movie_correctness` to a table with three columns, `'Title'`, `'Genre'`, and `'Was correct'`.

- The `'Genre'` column should contain the original genres, not the ones you predicted.
- The `'Was correct'` column should contain `True` or `False` depending on whether or not the movie was classified correctly.

```python
# Feel free to use multiple lines of code
# but make sure to assign test_movie_correctness to the proper table!
test_movie_correctness = ...
print(test_movie_correctness.head(5))
```

**Question 3.3.3 (Manually Graded)**

Do you see a pattern in the types of movies your classifier misclassifies? In two sentences or less, describe any patterns you see in the results or any other interesting findings from the table above. If you need some help, try looking up the movies that your classifier got wrong on Wikipedia.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## Working with all features (Using Sklearn)

Let's deep dive to fit the classficiation knn and random forest model on all predictors of the movies data set.

```python
X = movies.drop(columns=movies.columns[np.arange(5)], axis=1)
y = movies["Genre"]
```

**Question 4.1.1**

Complete the following Python class (named `KNN`) as per the docstrings mentioned to fit, predict, and evaluate the model's score. The `fit()` must have an option to perform hyperparameter tuning. The optimal values should be used such that the model accuracy must be more than `80%`.


```python
class KNN:
    def __init__(self, X, y, n_neighbors=5, test_size=0.2, random_state=42):
        """
        Initialize the KNN classifier.

        Parameters:
        -----------
            X (pandas.DataFrame): The feature matrix of shape (n_samples, n_features).
            y (pandas.Series): The target vector of shape (n_samples,).
            n_neighbors (int, optional): The number of nearest neighbors to use in classification. Defaults to 5.
            test_size (float, optional): The proportion of samples to use for testing. Defaults to 0.2.
            random_state (int, optional): The random state to use for splitting the data. Defaults to 42.
        """
        self.X_train, self.X_test, self.y_train, self.y_test = ...
        self.model = ...
        self.best_params = ...
        self.best_score = ...

    def fit(self, tune_fit="no"):
        """
        Fit the KNN model to the training data.

        Parameters:
        -----------
            tune_fit (str, optional): Whether to perform hyperparameter tuning. If "yes", performs a grid search to find
                the best hyperparameters. Defaults to "no".

        Raises:
        -----------
            ValueError: If `tune_fit` is not "yes" or "no".

        Returns:
        -----------
            None
        """
        if tune_fit=="yes":
            param_grid = ...
                ...
                ...
            }
            grid_search = ...
            ...
            self.best_params = ...
            self.model = ...
            self.best_score = ...
        elif tune_fit=="no":
            ...
        else:
            raise ValueError("Invalid value for `tune_fit`. Must be either 'yes' or 'no'.")

    def predict(self):
        """
        Predict the target values for the test data.

        Parameters:
        -----------
        None

        Returns:
        -----------
            numpy.ndarray: The predicted target values of shape (n_samples,)
        """
        ...

    def score(self):
        """
        Calculate the accuracy score and classification report for the KNN model.

        Parameters:
        -----------
        None

        Returns:
        -----------
        accuracy score: The ratio of the correctly predicted observations to the total observations.

        classification report: A text report of the main classification metrics such as precision, recall, f1-score and support for each class.
        """

        y_pred = ...
        ...

```

**Question 4.1.2**

Use the `KNN` class to fit, predict and report the accuracy, precision and recall of the model.

```python
fit_knn = ...
...
pred_knn = ...
accuracy_knn, classification_report_knn = ...
print("The accuracy achieved by knn model:", accuracy_knn)
print('#'*60)
print("The classification report of knn model: \n", classification_report_knn)
```

**Question 4.2.1**

Complete the following Python class (named `RandomForest`) as per the docstrings mentioned to fit, predict, and evaluate the model's score. The `fit()` must have an option to perform hyperparameter tuning. The optimal values should be used such that the model accuracy must be more than `83%`.


```python
class RandomForest:
    def __init__(self, X, y, random_state=RANDOM_STATE):
        """Initializes the RandomForest class with the input features X and target variable y,
        and the random state used for reproducibility.

        Parameters:
        -----------
        X : pandas.DataFrame
            The input feature matrix of shape (n_samples, n_features).
        y : pandas.Series
            The target variable of shape (n_samples,).
        random_state : int, default=RANDOM_STATE, which is 42
            The seed value for random number generator used to split the data.

        Returns:
        --------
        None"""

        self.X = ...
        self.y = ...
        self.random_state = ...
        self.model = ...
        self.best_params = {}
        self.best_score = 0
        self.X_train, self.X_test, self.y_train, self.y_test = train_test_split(self.X, self.y, test_size=0.2, random_state=random_state)

    def fit(self, tune_fit="yes"):
        """Trains the random forest model using the input data X and y.
        If tune_fit is set to "yes", it tunes the hyperparameters using GridSearchCV(), otherwise, with default parameters.
        It then stores the best hyperparameters and best estimator in the attributes best_params and model, respectively.

        Parameters:
        -----------
        tune_fit : str, default="yes"
            If "yes", tune the hyperparameters using GridSearchCV(), otherwise use default parameters.

        Returns:
        --------
        None
        """

        if tune_fit=="yes":
            param_grid = ...
            grid_search = ...
            ...
            self.best_params = ...
            self.model = ...
        else:
            ...

    def predict(self):
        """Predicts the target variable of the test data using the trained random forest model.
        Returns the predicted target variable values.

        Parameters:
        -----------
        None

        Returns:
        --------
        numpy.ndarray: The predicted target variable values

        """

        ...

    def score(self):
        """Computes the accuracy score and classification report for the predicted target variable and the actual test target variable.

        Parameters:
        -----------
        None

        Returns:
        --------
        accuracy score: The ratio of the correctly predicted observations to the total observations.

        classification report: A text report of the main classification metrics such as precision, recall, f1-score and support for each class."""

        y_pred = ...
        ...

    def plot_decision_tree(self):
        """Plots the first decision tree in the trained random forest model.
        Parameters:
        -----------
        None

        Returns:
        --------
        None
        """

        # Extract the first tree from the forest
        estimator = self.model.estimators_[0]

        # Plot the decision tree
        plt.figure(figsize=(12, 8))
        ...
        plt.draw()

    def plot_feature_importance(self, top=10):
        """Plots a horizontal bar chart of the top (by default 10) important features in the random forest model.

        Parameters:
        -----------
        top: The number of top important features to display. Default is 10.

        Returns:
        -----------
        None"""
        importances = ...
        indices = ...
        plt.title("Feature Importance")
        ...
        plt.yticks(range(top), [self.X.columns[i] for i in indices[-top:]])
        plt.xlabel("Relative Importance")
        plt.draw()

```

**Question 4.2.2**

Use the `RandomForest` class to fit, predict and report the accuracy, precision and recall of the model on movies dataset (X, y).

```python
fit_rf = ...
...
...
accuracy_rf, classification_report_rf = ...
print("The accuracy achieved by random forest model:", accuracy_rf)
print('#'*60)
print("The classification report of random forest model: \n", classification_report_rf)
fit_rf.plot_decision_tree()
fit_rf.plot_feature_importance(top=10)
```

**Question 4.3.1 (Maually Graded)**

1. Which model is better in terms of accuracy?
2. Which model has the highest recall score and why?
3. What can you infer about the performance of the two models based on their precision and recall scores?
4. How do the precision, recall, and F1 scores change when the hyper-parameters of the KNN and Random Forest models are tuned?
5. How do the precision, recall, and F1 scores vary when the KNN model is trained on different values of k or when the Random Forest model is trained on a different number of decision trees?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

**Question 4.3.2 (Maually Graded)**

1. How did you tune the hyperparameters of the KNN and Random Forest algorithms, and what was the effect of hyperparameter tuning on the performance of the classification algorithm?

2. Can you explain the decision boundary of the classification algorithm for the movie genre classification problem? How does the decision boundary change with different values of k in the KNN algorithm, or with different number of decision trees in the Random Forest algorithm?
   
3. If the data used to train the KNN and Random Forest models was imbalanced such that one class had significantly fewer samples than the others, which model would be more effective at correctly identifying samples from the underrepresented class based on the recall score, and why?

4. How would the precision, recall, and F1 scores for the KNN and Random Forest models change if a new class was added to the classification problem, and why?

5. How does the number of decision trees used in the Random Forest model impact the precision and recall scores, and what is the optimal number of decision trees to use based on the given classification report?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## <center><u> End of the Assignment 3 </u> </center>
### <center>Thanks</center>

Tests for q1_0 (points: [0])
```
>>> 0 < expected_row_sum
True
```


Tests for q1_1_1 (points: [0, 0])
```
>>> type(stemmed_message) == str
True
```
```
>>> len(stemmed_message) < len('elements')
True
```


Tests for q1_1_2 (points: [0, 1])
```
>>> type(most_stem) == str
True
```
```
>>> most_stem == 'gener'
True
```


Tests for q1_1_3 (points: [1])
```
>>> type(longest_uncut) == str
True
```


Tests for q1_2_1 (points: [0])
```
>>> 0.2 < outer_space_cor < 0.4
np.True_
```


Tests for q1_2_3 (points: [0])
```
>>> type(san_francisco) == int
True
```


Tests for q2_1_1 (points: [0])
```
>>> 0 < one_distance < 0.01
np.True_
```


Tests for q2_1_2 (points: [1, 2])
```
>>> correct_dis = 0.001406116
>>> dis = distance_two_features('clerks.', 'the godfather', 'water', 'feel')
>>> np.isclose(np.round(dis, 9), correct_dis)
np.True_
```
```
>>> correct_dis = 0.005647119
>>> dis = distance_two_features('clerks.', 'the godfather', 'your', 'that')
>>> np.isclose(np.round(dis, 9), correct_dis)
np.True_
```


Tests for q2_1_3 (points: [0])
```
>>> np.isclose(distance_from_python('clerks.'), 0.00079838)
np.True_
```


Tests for q2_1_4 (points: [1, 0, 1])
```
>>> set(close_movies.columns) >= {'Genre', 'Title', 'feel', 'water'}
True
```
```
>>> close_movies.shape[0] == 5
True
```
```
>>> close_movies['Title'].iloc[0] != 'monty python and the holy grail'
True
```


Tests for q2_1_5 (points: [3])
```
>>> most_common(close_movies, 'Genre') == 'thriller'
True
```


Tests for q3_0 (points: [0.5, 0.5, 0])
```
>>> 0.0 <= distance_first_to_second <= 0.1
np.True_
```
```
>>> np.isclose(distance(np.array([1, 2]), np.array([1, 2])), 0)
np.True_
```
```
>>> np.isclose(distance(np.array([1, 2, 3]), np.array([2, 4, 5])), 3)
np.True_
```


Tests for q3_1_1 (points: [0])
```
>>> bottom_left >= 1 and bottom_left <= 5
True
```


Tests for q3_1_2 (points: [0])
```
>>> bottom_right >= 1 and bottom_right <= 5
True
```


Tests for q3_1_3 (points: [0])
```
>>> top_right >= 1 and top_right <= 5
True
```


Tests for q3_1_4 (points: [0])
```
>>> top_left >= 1 and top_left <= 5
True
```


Tests for q3_1_5 (points: [0])
```
>>> movie_genre_guess >= 1 and movie_genre_guess <= 2
True
```


Tests for q3_1_6 (points: [0, 0, 0, 1.5, 1.5])
```
>>> len(my_features) >= 10
True
```
```
>>> np.all([f in test_movies.columns for f in my_features])
np.True_
```
```
>>> len(set(my_features)) >= 10
True
```
```
>>> help_message = "It looks like there are many movies in the training set that \n                  don't have any of your chosen words.  That will make your\n                  classifier perform very poorly in some cases.  Try choosing\n                  at least 1 common word.\n                  "
>>> train_f = train_movies[my_features]
>>> np.count_nonzero(train_f.apply(lambda r: np.sum(np.abs(np.array(list(r)))) == 0)) < len(my_features)
np.True_
```
```
>>> help_message = "It looks like there are many movies in the test set that \n                  don't have any of your chosen words.  That will make your\n                  classifier perform very poorly in some cases.  Try choosing\n                  at least 1 common word.\n                  "
>>> test_f = test_movies[my_features]
>>> np.count_nonzero(test_f.apply(lambda r: np.sum(np.abs(np.array(list(r)))) == 0)) < 5
np.True_
```


Tests for q3_1_7 (points: [0, 1, 2])
```
>>> (genre_and_distances.columns == ['Genre', 'Distance']).all()
np.True_
```
```
>>> genre_and_distances.shape[0] == 283
True
```
```
>>> np.allclose(genre_and_distances['Distance'], sorted(fast_distances(test_my_features.iloc[0], train_my_features)))
True
```


Tests for q3_1_8 (points: [1, 1])
```
>>> genre_and_distances.take(np.arange(7)).groupby('Genre').count().loc[my_assigned_genre][0] >= 4
np.True_
```
```
>>> my_assigned_genre == 'comedy'
True
```


Tests for q3_2_1 (points: [3])
```
>>> from collections import Counter
>>> g = train_movies['Genre']
>>> def check(r, k):
...     t = test_my_features.iloc[r]
...     return classify(t, train_my_features, g, k) == most_common(pd.DataFrame({'Genre': g, 'Distance': fast_distances(t, train_my_features)}).sort_values('Distance')[:k], 'Genre')
>>> check_5_nn = [check(i, 5) for i in np.arange(11)]
>>> all(check_5_nn)
True
```


Tests for q3_2_2 (points: [3])
```
>>> godzilla_genre == 'thriller'
True
```


Tests for q3_2_3 (points: [3])
```
>>> classify(test_my_features.iloc[0], train_my_features, train_movies['Genre'], 10) == 'comedy'
True
```


Tests for q3_3_1 (points: [0])
```
>>> 0 <= proportion_correct <= 1
np.True_
```


Tests for q3_3_2 (points: [2])
```
>>> assert test_movie_correctness.shape[0] == test_movies.shape[0]
```


Tests for q4_1_1 (points: [10])
```
>>> fit_knn_t = KNN(X, y)
>>> fit_knn_t.fit(tune_fit='yes')
>>> pred_knn = fit_knn_t.predict()
>>> assert fit_knn_t.model is not None
>>> assert fit_knn_t.best_params != {}
>>> assert fit_knn_t.best_score > 0
>>> assert isinstance(pred_knn, np.ndarray)
>>> assert len(pred_knn) == len(fit_knn_t.y_test)
>>> accuracy, report = fit_knn_t.score()
>>> assert isinstance(report, str)
>>> assert accuracy >= 0 and accuracy <= 1
```


Tests for q4_1_2 (points: [0, 3])
```
>>> accuracy_knn >= 0.82
True
```
```
>>> fit_knn is not None
True
```


Tests for q4_2_1 (points: [2, 2, 2, 2, 2])
```
>>> fit_rf_t = RandomForest(X, y)
>>> fit_rf_t.fit(tune_fit='no')
>>> pred_rf = fit_rf_t.predict()
>>> assert fit_rf_t.model is not None
```
```
>>> fit_rf_t = RandomForest(X, y)
>>> fit_rf_t.fit(tune_fit='no')
>>> pred_rf = fit_rf_t.predict()
>>> assert isinstance(pred_rf, np.ndarray)
```
```
>>> fit_rf_t = RandomForest(X, y)
>>> fit_rf_t.fit(tune_fit='no')
>>> pred_rf = fit_rf_t.predict()
>>> assert len(pred_rf) == len(fit_rf_t.y_test)
```
```
>>> fit_rf_t = RandomForest(X, y)
>>> fit_rf_t.fit(tune_fit='no')
>>> pred_rf = fit_rf_t.predict()
>>> accuracy, report = fit_rf_t.score()
>>> assert isinstance(report, str)
```
```
>>> fit_rf_t = RandomForest(X, y)
>>> fit_rf_t.fit(tune_fit='no')
>>> pred_rf = fit_rf_t.predict()
>>> accuracy, report = fit_rf_t.score()
>>> assert accuracy >= 0 and accuracy <= 1
```


Tests for q4_2_2 (points: [3, 0])
```
>>> accuracy_rf >= 0.83
True
```
```
>>> fit_rf is not None
True
```
